<problem display_name="Question with Partial Credit Possible" markdown="null">
  <text>
    <p>
      <h4>How Big Is Space?</h4>
    </p>
    <script type="text/python" system_path="python_lib">

def test_str(expect, ans):
  print expect, ans
  ans = ans.strip("'")
  ans = ans.strip('"')
  ans = float(ans)
  expect = float(expect)
  if abs(ans-expect) &amp;lt; 5:
    return True
  else:
    return False
def hint_fn(answer_ids, student_answers, new_cmap, old_cmap):
  aid = answer_ids[0]
  ans = str(student_answers[aid])
  print 'hint_fn called, ans=', ans
  hint = ''
  if float(ans) &lt; 51:
    hint = 'Much too low.'
    newgrade = 0.2
  elif float(ans) &lt;= 509:
    hint = 'Too low.'
    newgrade = 0.5
  elif float(ans) &gt; 5140:
    hint = 'Much too high.'
    newgrade = 0.2
  elif float(ans) &gt;= 519:
    hint = 'Too high.'
    newgrade = 0.5
  else:
    hint = 'Exact answer: 514 miles.'
    newgrade = 1
  if hint:
    hint = "&lt;font color='blue'&gt;Hint: {0}&lt;/font&gt;".format(hint)
    new_cmap.set_hint_and_mode(aid,hint,'always')
    new_cmap.set_property(aid,'npoints',newgrade)

</script>
    <p>
Imagine a scaled-down universe. If the sun is at one end of a football field and Neptune is at the other, how many miles away will the nearest star be? </p>
    <p>
      <customresponse cfn="test_str" expect="514">
        <textline correct_answer="514"/>
        <hintgroup hintfn="hint_fn"/>
      </customresponse>
    </p>
  </text>
</problem>
